rule_files:
  - "codeflare_alerts.yaml"

evaluation_interval: 1m

tests:
  # burn rate
  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="codeflare-operator"}
        values: "0x60"
      - series: probe_success:burnrate30m{instance="codeflare-operator"}
        values: "0x60"
      - series: probe_success:burnrate1h{instance="codeflare-operator"}
        values: "0x60"
      - series: probe_success:burnrate2h{instance="codeflare-operator"}
        values: "0x60"
      - series: probe_success:burnrate6h{instance="codeflare-operator"}
        values: "0x60"
      - series: probe_success:burnrate1d{instance="codeflare-operator"}
        values: "0x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: CodeFlare Operator Probe Success Burn Rate
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="codeflare-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate1h{instance="codeflare-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 2m
        alertname: CodeFlare Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: CodeFlare Operator Probe Success Burn Rate
              instance: "codeflare-operator"
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              message: "High error budget burn for codeflare-operator (current value: 3)."
              summary: CodeFlare Operator Probe Success Burn Rate
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Distributed-Workloads/codeflare-operator-availability.md'

  - interval: 1m
    input_series:
      - series: probe_success:burnrate30m{instance="codeflare-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate6h{instance="codeflare-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 15m
        alertname: CodeFlare Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: CodeFlare Operator Probe Success Burn Rate
              instance: "codeflare-operator"
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              message: "High error budget burn for codeflare-operator (current value: 16)."
              summary: CodeFlare Operator Probe Success Burn Rate
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Distributed-Workloads/codeflare-operator-probe-success-burn-rate.md'

  - interval: 1m
    input_series:
      - series: probe_success:burnrate2h{instance="codeflare-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate1d{instance="codeflare-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: CodeFlare Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: CodeFlare Operator Probe Success Burn Rate
              instance: "codeflare-operator"
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              message: "High error budget burn for codeflare-operator (current value: 61)."
              summary: CodeFlare Operator Probe Success Burn Rate
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Distributed-Workloads/codeflare-operator-probe-success-burn-rate.md'

  # operator running
  - interval: 1m
    input_series:
      - series: up{job="CodeFlare Operator"}
        values: 1
    alert_rule_test:
      - eval_time: 1m
        alertname: CodeFlare Operator is not running
        exp_alerts: []
  
  - interval: 1m
    input_series:
      - series: up{job="CodeFlare Operator"}
        values: 0
    alert_rule_test:
      - eval_time: 1m
        alertname: CodeFlare Operator is not running
        exp_alerts:
          - exp_labels:
              alertname: CodeFlare Operator is not running
              job: "CodeFlare Operator"
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              description: This alert fires when the CodeFlare Operator is not running.
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Distributed-Workloads/codeflare-operator-availability.md'
              summary: Alerting for CodeFlare Operator
  
  - interval: 1m
    input_series:
    alert_rule_test:
      - eval_time: 2m
        alertname: CodeFlare Operator taking too long to be up
        exp_alerts:
          - exp_labels:
              alertname: CodeFlare Operator taking too long to be up
              namespace: "redhat-ods-applications"
              job: "CodeFlare Operator"
              severity: info
            exp_annotations:
              description: This alert fires when the CodeFlare Operator takes over 2 min. to come back online. Either CodeFlare Operator is not running and failing to become ready, is misconfigured, or the metrics endpoint is not responding.
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Distributed-Workloads/codeflare-operator-absent-over-time.md'
              summary: Alerting for CodeFlare Operator
