rule_files:
  - data_science_pipelines_operator_alerts.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total:burnrate2m{component="dsp"}
        values: "0x60"
      - series: haproxy_backend_http_responses_total:burnrate5m{component="dsp"}
        values: "0x60"
      - series: haproxy_backend_http_responses_total:burnrate1h{component="dsp"}
        values: "0x60"
      - series: haproxy_backend_http_responses_total:burnrate2h{component="dsp"}
        values: "0x60"
      - series: haproxy_backend_http_responses_total:burnrate6h{component="dsp"}
        values: "0x60"
      - series: haproxy_backend_http_responses_total:burnrate1d{component="dsp"}
        values: "0x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: Data Science Pipelines Application Route Error Burn Rate
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total:burnrate5m{component="dsp"}
        values: "1+1x60"
      - series: haproxy_backend_http_responses_total:burnrate1h{component="dsp"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 2m
        alertname: Data Science Pipelines Application Route Error Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipelines Application Route Error Burn Rate
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              summary: "Data Science Pipelines Application Route Error Burn Rate"
              message: "High error budget burn for  (current value: 3)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-application-error-burn-rate.md'

  - interval: 1m
    input_series: 
      - series: haproxy_backend_http_responses_total:burnrate30m{component="dsp"}
        values: "1+1x60"
      - series: haproxy_backend_http_responses_total:burnrate6h{component="dsp"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 15m
        alertname: Data Science Pipelines Application Route Error Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipelines Application Route Error Burn Rate
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              summary: "Data Science Pipelines Application Route Error Burn Rate"
              message: "High error budget burn for  (current value: 16)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-application-error-burn-rate.md'

  - interval: 1m
    input_series: 
      - series: haproxy_backend_http_responses_total:burnrate2h{component="dsp"}
        values: "1+1x60"
      - series: haproxy_backend_http_responses_total:burnrate1d{component="dsp"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: Data Science Pipelines Application Route Error Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipelines Application Route Error Burn Rate
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              summary: "Data Science Pipelines Application Route Error Burn Rate"
              message: "High error budget burn for  (current value: 61)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-application-error-burn-rate.md'

  - interval: 1m
    input_series: 
      - series: haproxy_backend_http_responses_total:burnrate6h{component="dsp"}
        values: "1+1x200"
      - series: haproxy_backend_http_responses_total:burnrate3d{component="dsp"}
        values: "1+1x200"
    alert_rule_test:
      - eval_time: 3h
        alertname: Data Science Pipelines Application Route Error Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipelines Application Route Error Burn Rate
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              summary: "Data Science Pipelines Application Route Error Burn Rate"
              message: "High error budget burn for  (current value: 181)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-application-error-burn-rate.md'

  # Probe success
  - interval: 1m
    input_series:
      - series: probe_success:burnrate2m{instance="data-science-pipelines-operator"}
        values: "0x60"
      - series: probe_success:burnrate5m{instance="data-science-pipelines-operator"}
        values: "0x60"
      - series: probe_success:burnrate1h{instance="data-science-pipelines-operator"}
        values: "0x60"
      - series: probe_success:burnrate2h{instance="data-science-pipelines-operator"}
        values: "0x60"
      - series: probe_success:burnrate6h{instance="data-science-pipelines-operator"}
        values: "0x60"
      - series: probe_success:burnrate1d{instance="data-science-pipelines-operator"}
        values: "0x60"
    alert_rule_test:
      - eval_time: 3h
        alertname: Data Science Pipelines Operator Probe Success Burn Rate
        exp_alerts: []


  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="data-science-pipelines-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate1h{instance="data-science-pipelines-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 2m
        alertname: Data Science Pipelines Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipelines Operator Probe Success Burn Rate
              instance: "data-science-pipelines-operator"
              namespace: "redhat-ods-applications"
              severity: critical
            exp_annotations:
              summary: "Data Science Pipelines Operator Probe Success Burn Rate"
              message: "High error budget burn for data-science-pipelines-operator (current value: 3)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-operator-probe-success-burn-rate.md'

  - interval: 1m
    input_series: 
      - series: probe_success:burnrate30m{instance="data-science-pipelines-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate6h{instance="data-science-pipelines-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 15m
        alertname: Data Science Pipelines Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipelines Operator Probe Success Burn Rate
              instance: "data-science-pipelines-operator"
              namespace: "redhat-ods-applications"
              severity: critical
            exp_annotations:
              summary: "Data Science Pipelines Operator Probe Success Burn Rate"
              message: "High error budget burn for data-science-pipelines-operator (current value: 16)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-operator-probe-success-burn-rate.md'

  - interval: 1m
    input_series: 
      - series: probe_success:burnrate2h{instance="data-science-pipelines-operator"}
        values: "1+1x60"
      - series: probe_success:burnrate1d{instance="data-science-pipelines-operator"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: Data Science Pipelines Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipelines Operator Probe Success Burn Rate
              instance: "data-science-pipelines-operator"
              namespace: "redhat-ods-applications"
              severity: warning
            exp_annotations:
              summary: "Data Science Pipelines Operator Probe Success Burn Rate"
              message: "High error budget burn for data-science-pipelines-operator (current value: 61)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-operator-probe-success-burn-rate.md'

  # application unavailable
  - interval: 1m
    input_series: 
      - series: data_science_pipelines_application_ready{dspa_name="dspa_instance_1", dspa_namespace="dspa_namespace_a"}
        values: "0x200"
      - series: data_science_pipelines_application_apiserver_ready{dspa_name="dspa_instance_1", dspa_namespace="dspa_namespace_a"}
        values: "0x200"
      - series: data_science_pipelines_application_persistenceagent_ready{dspa_name="dspa_instance_1", dspa_namespace="dspa_namespace_a"}
        values: "0x200"
      - series: data_science_pipelines_application_scheduledworkflow_ready{dspa_name="dspa_instance_1", dspa_namespace="dspa_namespace_a"}
        values: "0x200"
    alert_rule_test:
      - eval_time: 2m
        alertname: Data Science Pipeline Application Unavailable
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipeline Application Unavailable
              dspa_name: "dspa_instance_1"
              dspa_namespace: "dspa_namespace_a"
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              message: "Data Science Pipelines Application is down!"
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-application-component-readiness-status.md'
              summary: The Data Science Pipelines Application CustomResource "dspa_instance_1" in namespace "dspa_namespace_a" has been NotReady for more than 5 minutes
      - eval_time: 2m
        alertname: Data Science Pipeline APIServer Unavailable
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipeline APIServer Unavailable
              dspa_name: "dspa_instance_1"
              dspa_namespace: "dspa_namespace_a"
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              message: "Data Science Pipelines APIServer component is down!"
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-application-component-readiness-status.md'
              summary: A Data Science Pipelines APIServer pod owned by DSPA "dspa_instance_1" in namespace "dspa_namespace_a" has been NotReady for more than 5 minutes
      - eval_time: 2m
        alertname: Data Science Pipeline PersistenceAgent Unavailable
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipeline PersistenceAgent Unavailable
              dspa_name: "dspa_instance_1"
              dspa_namespace: "dspa_namespace_a"
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              message: "Data Science Pipelines PersistenceAgent component is down!"
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-application-component-readiness-status.md'
              summary: A Data Science Pipelines PersistenceAgent pod owned by DSPA "dspa_instance_1" in namespace "dspa_namespace_a" has been NotReady for more than 5 minutes
      - eval_time: 2m
        alertname: Data Science Pipeline ScheduledWorkflows Unavailable
        exp_alerts:
          - exp_labels:
              alertname: Data Science Pipeline ScheduledWorkflows Unavailable
              dspa_name: "dspa_instance_1"
              dspa_namespace: "dspa_namespace_a"
              namespace: "redhat-ods-applications"
              severity: info
            exp_annotations:
              message: "Data Science Pipelines ScheduledWorkflows component is down!"
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Data-Science-Pipelines/data-science-pipelines-application-component-readiness-status.md'
              summary: A Data Science Pipelines ScheduledWorkflow controller pod owned by DSPA "dspa_instance_1" in namespace "dspa_namespace_a" has been NotReady for more than 5 minutes
