rule_files:
  - feast-operator-alerting.rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="feast-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate1h{instance="feast-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate30m{instance="feast-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate6h{instance="feast-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate2h{instance="feast-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate1d{instance="feast-operator-controller-manager"}
        values: "0x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: Feast Operator Probe Success Burn Rate
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="feast-operator-controller-manager"}
        values: "1+1x60"
      - series: probe_success:burnrate1h{instance="feast-operator-controller-manager"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 2m
        alertname: Feast Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Feast Operator Probe Success Burn Rate
              instance: "feast-operator-controller-manager"
              severity: critical
            exp_annotations:
              summary: "Feast Operator Probe Success Burn Rate"
              message: "High error budget burn for feast-operator-controller-manager (current value: 3)."
              triage: "https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Feature-Store/rhoai-feast-operator-probe-success-burn-rate.md"

  - interval: 1m
    input_series:
      - series: probe_success:burnrate30m{instance="feast-operator-controller-manager"}
        values: "1+1x60"
      - series: probe_success:burnrate6h{instance="feast-operator-controller-manager"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 15m
        alertname: Feast Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Feast Operator Probe Success Burn Rate
              instance: "feast-operator-controller-manager"
              severity: critical
            exp_annotations:
              summary: "Feast Operator Probe Success Burn Rate"
              message: "High error budget burn for feast-operator-controller-manager (current value: 16)."
              triage: "https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Feature-Store/rhoai-feast-operator-probe-success-burn-rate.md"

  - interval: 1m
    input_series:
      - series: probe_success:burnrate2h{instance="feast-operator-controller-manager"}
        values: "1+1x60"
      - series: probe_success:burnrate1d{instance="feast-operator-controller-manager"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: Feast Operator Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: Feast Operator Probe Success Burn Rate
              instance: "feast-operator-controller-manager"
              severity: warning
            exp_annotations:
              summary: "Feast Operator Probe Success Burn Rate"
              message: "High error budget burn for feast-operator-controller-manager (current value: 61)."
              triage: "https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Feature-Store/rhoai-feast-operator-probe-success-burn-rate.md"
