rule_files:
  - model_controller_alerts.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="odh-model-controller"}
        values: "0x60"
      - series: probe_success:burnrate1h{instance="odh-model-controller"}
        values: "0x60"
      - series: probe_success:burnrate30m{instance="odh-model-controller"}
        values: "0x60"
      - series: probe_success:burnrate6h{instance="odh-model-controller"}
        values: "0x60"
      - series: probe_success:burnrate2h{instance="odh-model-controller"}
        values: "0x60"
      - series: probe_success:burnrate1d{instance="odh-model-controller"}
        values: "0x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: ODH Model Controller Probe Success Burn Rate
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="odh-model-controller"}
        values: "1+1x60"
      - series: probe_success:burnrate1h{instance="odh-model-controller"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 2m
        alertname: ODH Model Controller Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: ODH Model Controller Probe Success Burn Rate
              namespace: "redhat-ods-applications"
              instance: "odh-model-controller"
              severity: critical
            exp_annotations:
              summary: "ODH Model Controller Probe Success Burn Rate"
              message: "High error budget burn for odh-model-controller (current value: 3)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhods-odh-controller-probe-success-burn-rate.md'

  - interval: 1m
    input_series:
      - series: probe_success:burnrate30m{instance="odh-model-controller"}
        values: "1+1x60"
      - series: probe_success:burnrate6h{instance="odh-model-controller"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 15m
        alertname: ODH Model Controller Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: ODH Model Controller Probe Success Burn Rate
              namespace: "redhat-ods-applications"
              instance: "odh-model-controller"
              severity: critical
            exp_annotations:
              summary: "ODH Model Controller Probe Success Burn Rate"
              message: "High error budget burn for odh-model-controller (current value: 16)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhods-odh-controller-probe-success-burn-rate.md'
 
  - interval: 1m
    input_series:
      - series: probe_success:burnrate2h{instance="odh-model-controller"}
        values: "1+1x60"
      - series: probe_success:burnrate1d{instance="odh-model-controller"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: ODH Model Controller Probe Success Burn Rate
        exp_alerts:
          - exp_labels:
              alertname: ODH Model Controller Probe Success Burn Rate
              namespace: "redhat-ods-applications"
              instance: "odh-model-controller"
              severity: warning
            exp_annotations:
              summary: "ODH Model Controller Probe Success Burn Rate"
              message: "High error budget burn for odh-model-controller (current value: 61)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhods-odh-controller-probe-success-burn-rate.md'
