rule_files:
  - operator-alerting.rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: kube_pod_container_status_restarts_total{namespace="redhat-ods-operator", pod="rhods-operator-abc123", container="manager"}
        values: "0 0 0 0 0 0"
    alert_rule_test:
      - eval_time: 7m
        alertname: OperatorPodRestartingFrequently
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: kube_pod_container_status_restarts_total{namespace="redhat-ods-operator", pod="rhods-operator-abc123", container="manager"}
        values: "0 1 2 3 3 3 3 3 3"
    alert_rule_test:
      - eval_time: 7m
        alertname: OperatorPodRestartingFrequently
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: kube_pod_container_status_restarts_total{namespace="redhat-ods-operator", pod="rhods-operator-abc123", container="manager"}
        values: "0 0 0 1 2 3 4 5 6 7 8 9 10"
    alert_rule_test:
      - eval_time: 9m
        alertname: OperatorPodRestartingFrequently
        exp_alerts:
          - exp_labels:
              alertname: OperatorPodRestartingFrequently
              severity: warning
              namespace: redhat-ods-operator
              pod: rhods-operator-abc123
              container: manager
            exp_annotations:
              summary: "Operator pod rhods-operator-abc123 is restarting frequently"
              description: "Pod rhods-operator-abc123 in namespace redhat-ods-operator has restarted 5 times in the last 5 minutes. This indicates potential instability."
              message: "Container manager in pod rhods-operator-abc123 (namespace: redhat-ods-operator) has been restarted more than 3 times in a 5-minute period."
              triage: "https://github.com/opendatahub-io/opendatahub-operator/blob/main/docs/troubleshooting.md#operator-pod-restarting-frequently"

  # Alert should NOT fire for pods in different namespace
  - interval: 1m
    input_series:
      - series: kube_pod_container_status_restarts_total{namespace="different-namespace", pod="some-operator-abc", container="manager"}
        values: "0 0 1 3 5 7 9 11 13"
    alert_rule_test:
      - eval_time: 7m
        alertname: OperatorPodRestartingFrequently
        exp_alerts: []

